<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 10  -->
  <question type="coderunner">
    <name>
      <text>C get current working directory</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Write a function <code>get_current_dir</code> that returns the absolute path to the current working directory. The returned pointer must point to a valid string. You may use <code>getcwd</code> from <code>&lt;unistd.h&gt;</code>. Do not print anything from the function.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_function</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>30</answerboxlines>
    <answerboxcolumns>200</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[
#include <stdio.h>
#include <stdlib.h>
#include <unistd.h>

{{ STUDENT_ANSWER }}

int main() {
    char *cwd = get_current_dir();
    printf("%s\n", cwd);
    free(cwd);
    return 0;
}
]]></template>
    <iscombinatortemplate>0</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[
char *get_current_dir() {
    char *buf = malloc(1024);
    if (!buf) return NULL;
    if (getcwd(buf, 1024) == NULL) {
        free(buf);
        return NULL;
    }
    return buf;
}
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang>sh</acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams>{"CFLAGS": "-Wall -Werror"}</sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
        <testcode>
                <text><![CDATA[
system("mkdir -p test_dir && cd test_dir");
chdir("test_dir");
char *cwd = get_current_dir();
printf("%s\n", cwd);
free(cwd);
]]></text>
        </testcode>
        <stdin>
                <text></text>
        </stdin>
        <expected>
                <text><![CDATA[/tmp/test_dir]]></text>
        </expected>
        <extra>
                <text></text>
        </extra>
        <display>
                <text>SHOW</text>
        </display>
      </testcase>
    </testcases>
  </question>
</quiz>