<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 11  -->
  <question type="coderunner">
    <name>
      <text>C program get CPU core count</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Write a complete C program that prints the number of available CPU cores on a Linux system.
Use <code>sysconf(_SC_NPROCESSORS_ONLN)</code> to get the number of cores.
Your program must compile and run without errors or warnings.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>30</answerboxlines>
    <answerboxcolumns>200</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[
#include <stdio.h>
#include <unistd.h>

{{ STUDENT_ANSWER }}
]]></template>
    <iscombinatortemplate>0</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[
int main(){
    int cpus = sysconf(_SC_NPROCESSORS_ONLN);
    printf("%d\n", cpus);
    return 0;
}
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang>sh</acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams>{"CFLAGS": "-Wall -Werror"}</sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
        <testcode>
                <text></text>
        </testcode>
        <stdin>
                <text></text>
        </stdin>
        <expected>
                <text>8</text>
        </expected>
        <extra>
                <text></text>
        </extra>
        <display>
                <text>SHOW</text>
        </display>
      </testcase>
    </testcases>
  </question>
</quiz>