<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 17  -->
  <question type="coderunner">
    <name>
      <text>C fork a child process</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Write a function <code>int spawn_process()</code> that creates a child process using <code>fork()</code>. The function must:</p>
<ul>
  <li>Return 1 in the <strong>child process</strong></li>
  <li>Return 0 in the <strong>parent process</strong></li>
  <li>Return -1 if <code>fork()</code> fails</li>
</ul>
<p>The test will verify that both the child and parent return correctly. You do not need to print the process ID.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_function</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>10, 20, ...</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>10</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template><![CDATA[
#include <stdio.h>
#include <stdlib.h>
#include <sys/types.h>
#include <unistd.h>

{{ STUDENT_ANSWER }}

int main() {
    int res = spawn_process();
    if (res == -1) {
        printf("fail\n");
    } else if (res == 1) {
        printf("child\n");
    } else if (res == 0) {
        printf("parent\n");
    } else {
        printf("invalid\n");
    }
    return 0;
}
]]></template>
    <iscombinatortemplate>0</iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[
int spawn_process() {
    pid_t pid = fork();
    if (pid == -1) {
        return -1; // Error
    } else if (pid == 0) {
        return 1; // Child
    } else {
        return 0; // Parent
    }
}
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang>sh</acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams>{"CFLAGS": "-Wall -Werror"}</sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>child
parent</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>